# Lab 2: Numpy
In this Numpy exercise, the general requirement is not to use loops; I will specify where it is allowed

(Last update: 12/11/2023)

Name: ...  
Sdudent ID: ...

---


## 0. Instructions for doing and submitting assignment

**How to do your assignment**

You will do your assignment directly on this notebook file. First, you fill your name and student code at the beginning of the file. In this file, you will write your code when you see the following lines of code:
```python
# YOUR CODE HERE
raise NotImplementedError()
```

For optional coding parts, there will be:
```python
# YOUR CODE HERE (OPTION)
```

For markdown cell, there will be:
```markdown
YOUR ANSWER HERE
```

Of course, you have to remove the `raise NotImplementedError()` statement when you finish.

For coding parts, there are often cells below to help you check your answers. You will pass the test if there are no errors when you run the test cells. In some cases, the tests are insufficient. That means if you do not pass the test, your answer is definitely wrong somewhere, but if you pass the test, your answer may still be incorrect.

While doing the assignment, you should print out the output and create more cells for testing. But you have to remove all of them (comment your print-out codes, delete the cell created by you) when you submit your code. <font color=red>Do not remove or edit my cells</font> (except for the aforementioned cells).

Keep your code clean and clear by using meaningful variable names and comments, not write too-long coding lines.
Press `Ctrl + S` right after editing.

Keep it real: The reason why you are here is to <font color=green>study, really study</font>. I highly recommend that you discuss your idea with your friends and <font color=green>write your own code based on your own knowledge</font>. <font color=red>Copy means zero.</font>

**How to submit your assignment**

When grading your assignment, I will choose `Kernel` - `Restart & Run All` in order to restart the kernel and run all cells in your notebook. Therefore, you should do that before submitting to ensure that the outputs are all as expected.

After that, rename the notebook as `<Student ID>.ipynb`. For example, if your student code is 1234567, then your notebook is `1234567.ipynb`.

Finally, submit your notebook file on Moodle. <font color=red>Please strictly follow the submission rules.</font>

---

## 1. Programming environment

- You will re-use the Linux environment setup in Lab 0 - WarmUp. Don't forget to start your coding environment (`conda activate min_ds-env`) before doing your assignment.
- Use Jupyter notebook or Jupyter lab, <font color=red>not Google Colab</font> (I can not grade you well on Google Colab) to edit your `*.ipynb` file.

```python
import sys
sys.executable
```

- If there are no problems, the file to run python will be the file of the "min_ds-env" code environment.

- In this article, you are not using the Pandas library.

---

## 2. Import necessary libraries

```python
import numpy as np
from zlib import adler32
# YOUR CODE HERE (OPTION)
```

---

## 3. Data collection

Numpy is not a great library for handling operations like data reading and writing, but it's an excellent library for computational tasks. Therefore, in this article, we'll simply use the pre-collected dataset that I've attached in the folder of this lab. This dataset actually contains multiple files and is relatively large, but it has been curated to include relevant information for this lab. You can learn more about this dataset [here](https://files.grouplens.org/datasets/movielens/ml-100k/u.data).

Here is a specific description of the dataset:
> MovieLens data sets were collected by the GroupLens Research Project
at the University of Minnesota.
> 
> This data set consists of:
> * 100,000 ratings (1-5) from 943 users on 1682 movies. 
> * Each user has rated at least 20 movies.  
> * Simple demographic info for the users (age, gender, occupation, zip)
>
> The data was collected through the MovieLens web site
(movielens.umn.edu) during the seven-month period from September 19th, 
1997 through April 22nd, 1998. This data has been cleaned up - users
who had less than 20 ratings or did not have complete demographic
information were removed from this data set. Detailed descriptions of
the data file can be found at the end of this file.

---

## 4. Data exploring & Data preprocessing

### How many rows and columns does the data have?

Of course, the first thing you need to do is read the data file into the Numpy array and name it `raw_ratings` (use function `np.genfromtxt`). You may encounter some minor problems with this task, it seems that all the data in the Numpy array is not what we want. This happens because the function `np.genfromtxt` has a default data type of `float`, you need to find a way to convert it to `uint64`. You should put the dataset file in the same directory as this notebook file to simplify when passing the file name to the function. Finally, you need to calculate the number of rows and columns for this dataset, these two values are stored in two variables `n_rows` and `n_cols` respectively.

<!-- grade_id: cell-bc45cdfd9571ab52, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dfaa826c3a1df79b, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert raw_ratings.dtype == np.uint64
assert adler32(str(raw_ratings.ndim).encode()) == 3342387
assert adler32(str(n_rows).encode()) == 66847010
assert adler32(str(n_cols).encode()) == 3473461
raw_ratings[:5] # Look at the first 5 lines
```

### Rows

#### The meaning of each row

Each row in the data set shows some information about a user's score for a movie.

#### Does the data have duplicate rows?

You will test this case and save the results to the `have_duplicated_rows` variable. This variable will have the value True if the data has duplicate lines and will have the value False otherwise.

<!-- grade_id: cell-204832200dec58f4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f94a96b9a7fa030a, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert have_duplicated_rows == False
```

Great, so there are no duplicate rows. Next we will explore the columns.

### Columns

#### The meaning of each column
- The first column shows the user id
- The second column shows the movie id
- The third column shows the score the user gave for the movie
- The fourth column shows the time the user gave the score (expressed in seconds from a benchmark)

#### What data type does each column currently have?

```python
raw_ratings.dtype
```

At first glance, it seems that all columns are numeric. But in my opinion, the two columns `user_id` and `moive_id` should be classified into categorical groups. The reason for this is because both `user_id` and `movie_id` are simply identifiers and do not necessarily have an arithmetic relationship between the columns. Of course, this is just an objective perspective and not true for all cases, but to make it easier to work, in this lab we will agree with the above thought.

#### For each column with numeric datatype, how are the values distributed?

First, we need to see how many missing values the numeric columns have. This mission is quite 'difficult' ^^ so I will do it for you.

```python
np.sum(np.isnan(raw_ratings[:, 2:]), axis=0)
```

Great, so all numeric columns don't have any missing values.

Now, your job is to calculate the min, Q1(25%), median, Q3(75%) and max of these numeric columns. You will need to use the `np.percentile` function to do this. Then, the all values of each column are saved respectively into 3 Numpy arrays namely `rate_col_profile`, `rate_date_col_profile`. These are two Numpy arrays (one-dimensional), where `rate_col_profile` has a `dtype` of float, and `rate_date_col_profile` has a `dtype` of `'datetime64[s]'`.

<!-- grade_id: cell-bf0638c7501e73f1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc39e5f303f573da, points: 1, grade: true, solution: false -->
```python
# TEST
assert adler32(str(rate_col_profile).encode()) == 444269344
assert adler32(str(rate_date_col_profile).encode()) == 4242871962
```

#### For each column with categorical datatype, how are the values distributed?

Just like with numeric columns, we need to see if two categorical columns have missing values? (This is difficult so let me do it for you :v )

```python
np.sum(np.isnan(raw_ratings[:, :2]), axis=0)
```

Your task is to, for each column, calculate a list of 5 numbers: the number of distinct values, the value that appears least with its corresponding count (total of 2 numbers), and the value that appears most with its corresponding count (total of 2 numbers). You should store the 2 lists calculated for 2 columns in two variables, namely `user_col_profile` and `movie_col_profile`. If multiple users rate the least number of movies, we will agree to choose the user with the smallest id. And vice versa, if many users rate the most movies, we will choose the user with the largest id.

<!-- grade_id: cell-32713eaf68686a2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-909a31870f2070a6, points: 1, grade: true, solution: false -->
```python
# assert adler32(str(user_col_profile).encode()) == 1375015361
# assert adler32(str(movie_col_profile).encode()) == 1325142473
assert user_col_profile == [943, # Có chừng này user
                        19,  # Đây là user rate ít movie nhất
                        20,  # và đó là chừng này movie
                        405, # Đây là user rate nhiều movie nhất
                        737] # và đó là chừng này movie
assert movie_col_profile == [1682,#Có chừng này movie
                         599, #Đây là movie được ít user rate nhất
                         1,   #và đó là chừng này user
                         50,  #Đây là movie được nhiều user rate nhất
                         583] #và đó là chừng này user
```

Incidentally, we need to check the maximum and minimum values of the two columns `user_id` and `movie_id`:

```python
print('User id  - min & max:', 
      raw_ratings[:, 0].min(), '&', raw_ratings[:, 0].max()) 
print('Movie id - min & max:', 
      raw_ratings[:, 1].min(), '&', raw_ratings[:, 1].max()) 
```

---

## 5. Question

The previous section was just to warm you up before diving into the main content of this lab. Now, we have a bit better understanding of the dataset. We will attempt to pose meaningful questions and find answers using the data.

One interesting question to ask is: *For each different user, is it possible to recommend movies that the user has never watched before?*

Finding an answer to this question can be beneficial for both users and movie streaming service providers:
- Users: Users may want to watch a movie, but with so many options available, they may not know which one to choose. It would be convenient for users if the system could suggest a list of movies that they are likely to enjoy.
- Movie Streaming Service Providers: If the system makes good recommendations, it's more likely that users will watch and enjoy the movies. This, in turn, means users will continue to pay for the service.

---

### Preprocessing

First, we need to decide which information to use in building the movie recommendation system. Obviously, the columns `user_id`, `moive_id`, and `rating` are essential to perform this task. As for the column `date`, this column can still have value in practice when building a recommendation model. However, for simplicity, we will temporarily set aside this column here.

Based on 3 columns, you need to create a 2D NumPy matrix named `ratings`. In this matrix, the number of rows represents the number of users, while the number of columns represents the number of movie. So, `ratings[i, j]` will represent the rating that `user_i` has given to `movie_j`. "For movie series that the user has not rated, the value will be 'NaN'."

<!-- grade_id: cell-880a7016c1a5a315, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-28591c055b9636de, points: 1, grade: true, solution: false -->
```python
assert ratings.shape == (943, 1682)
missing_ratios = np.mean(np.isnan(ratings))
assert missing_ratios.round(4) == .9370
```

### Analyze data to answer the question

It would be much simpler if we used algorithms supported by other libraries. However, the main goal of this lab is to help you practice using the Numpy library. Therefore, you will have the opportunity to build a simple movie recommendation system from scratch using the provided data, utilizing only Numpy library. Please remember that Numpy doesn't favor loops, so you are only allowed to use loops where I explicitly permit.

In my opinion, there are two fundamental tasks in a movie recommendation system:

- First, you need to predict the ratings for movie that a user hasn't reviewed or watched yet.
- Second, you need to provide recommendations to users based on the top-rated movies that have been predicted.

It seems that the second task will become much simpler if we can accomplish the first task. One of the simplest ways to tackle task 1 is by computing the similarity between users and using this similarity to make predictions. However, there are some considerations to keep in mind. It's not feasible to compute similarity between all users at once, as it might lead to memory issues (even if you have enough memory, my computer is quite limited in that regard :<). One way to address this issue is to process a group of users at a time, referred to as `a batch`. To keep it simple, let's stick with a `batch_size = 32`, which I believe is a reasonable value. You should try to make your code work with a single batch first and then extend it to process all batches.

```python
batch_size = 32
filled_ratings = np.empty_like(ratings)
```

"First, you will try with a batch corresponding to users with indices from `start` to `end`."

```python
start = 0
end = batch_size
```

Step 1: Calculate the `similarities` array to show the similarity between each user in the current batch with all users in the entire dataset. This array will have a size of `batch_size` x `n_users` (`n_users` is the total number of users in the dataset), where `similarities[i, j]` indicates the similarity between `user_i` and `user_j`. In the case where two users have no common rated movies (when running, you may see a warning 'RuntimeWarning: Mean of empty slice'), you set the similarity to 0.

<!-- grade_id: cell-ed3883d8af0fb097, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-45055a8d86566a14, points: 1, grade: true, solution: false -->
```python
# TEST
assert adler32(str(similarities.shape).encode()) == 136184227
assert adler32(str(similarities[:3, :3].round(1)).encode()) == 3499233691
```

Step 2: calculate the `weights` matrix. The array `weights` will have the size `batch_size` x `n_users` x `n_movies` (where `n_movies` is the total number of movies). About how to calculate `weights`, you can refer to file `example.ipynb`.

When running, you will see the warning "RuntimeWarning: invalid value encountered in true_divide"; This is because the users who rate a movie under consideration all have a similarity of 0 with a user under review, resulting in normalization to 0/0 and the result is difficult. This case means there is not enough information to predict the score and in this article, you should leave it as it is.

<!-- grade_id: cell-39cd467a8309e23d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-322c70e976509914, points: 1, grade: true, solution: false -->
```python
# TEST
assert weights.shape == (32, 943, 1682)
assert np.sum(np.isnan(weights)) == 31119
```

Step 3: For each user in the batch under consideration, calculate the score (for all movies) by multiplying the score of all users with the corresponding weight in the `weight` array; then write each user's scores down to one line in the `filled_ratings` array.

<!-- grade_id: cell-bee0a7ab7e39d6fd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d89ae1305fcd5db, points: 1, grade: true, solution: false -->
```python
# TEST
filled_batch = filled_ratings[start:end]
filled_nanvals = filled_batch[np.isnan(ratings[start:end])]
assert np.array_equal(filled_nanvals[:13].round(1), np.array([3.6, 3.5, 4. , 3.8, 3.5, 3.2, 3.4, 3.1, 3.2, 3.7, 3.7, 3.2, 4.3]))
assert np.array_equal(filled_nanvals[-13:].round(1), np.array([3., 1., 2., 3., 4., 3., 2., 3., 1., 3., 2., 3., 3.]))
```

Great ! So your code has run on a batch, now it's time for you to use the `for` loop to cycle through all the batches in the data set.

<!-- grade_id: cell-965580ca501d6d9c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a7fb93e0ee5d82f1, points: 3, grade: true, solution: false -->
```python
# TEST
filled_nanvals = filled_ratings[np.isnan(ratings)]
assert np.array_equal(filled_nanvals[:13].round(1), np.array([3.6, 3.5, 4. , 3.8, 3.5, 3.2, 3.4, 3.1, 3.2, 3.7, 3.7, 3.2, 4.3]))
assert np.array_equal(filled_nanvals[-13:].round(1), np.array([3., 0., 2., 3., 4., 3., 2., 3., 0., 0., 0., 3., 3.]))
```